# Instructions

1. Add your name below.
2. Complete each question. Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", and delete and `throw NotImplementedError()` lines.
3. Where applicable, run the test cases *below* each question to check your work.
4. Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All). You can also use the **Validate** button to run all test cases.
5. Turn in your homework by going to the main screen in JupyterHub, clicking the Assignments menu, and submitting.



```python
"""
Name: Xiaochun Liang
"""
```

Output:
```
'\nName: Xiaochun Liang\n'
```

```python

```

## Import and inspect the dataset

Before we do anything, we need to load and inspect the data.

Relevant documentation:
* [`pd.read_csv`](https://pandas.pydata.org/docs/reference/api/pandas.read_csv.html).

```python
import pandas as pd
```

<!-- grade_id: cell-fc5a1da3b774e210, points: 0, grade: false, solution: true -->
```python
data = None
# TODO: Read the data using pd.read_csv from the file /etc/data/employees.csv
# YOUR CODE HERE
data = pd.read_csv("employees.csv")
data
```

Output:
```
Age Attrition     BusinessTravel  DailyRate              Department  \
0      41       Yes      Travel_Rarely       1102                   Sales   
1      49        No  Travel_Frequently        279  Research & Development   
2      37       Yes      Travel_Rarely       1373  Research & Development   
3      33        No  Travel_Frequently       1392  Research & Development   
4      27        No      Travel_Rarely        591  Research & Development   
...   ...       ...                ...        ...                     ...   
1465   36        No  Travel_Frequently        884  Research & Development   
1466   39        No      Travel_Rarely        613  Research & Development   
1467   27        No      Travel_Rarely        155  Research & Development   
1468   49        No  Travel_Frequently       1023                   Sales   
1469   34        No      Travel_Rarely        628  Research & Development   

      DistanceFromHome  Education EducationField  EmployeeCount  \
0             
```

<!-- grade_id: cell-17e31e47dc8a9098, points: 1, grade: true, solution: false -->
```python
# Automated tests - make sure you pass these!
assert data.shape == (1470, 35)
```

## Data Splitting

Before we do anything, we should split our data into a train and test set.


**Instructions**: Split `data` into an `X` DataFrame containing the features and a `y` Series containing only the "Attrition" label.

Relevant documentation:
* [`DataFrame.drop`](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.drop.html).

<!-- grade_id: cell-3d02a7f4456d8874, points: 0, grade: false, solution: true -->
```python
# TODO: Create the X and y variables
# YOUR CODE HERE
X = data.drop(columns = ['Attrition'])
# y = data.get(['Attrition'])
y = data['Attrition']
```

<!-- grade_id: cell-e9fc8e72014d0118, points: 2, grade: true, solution: false -->
```python
assert 'Attrition' not in X.columns, "Make sure to drop the Attrition column"
assert X.shape == (1470, 34)

assert y.shape == (1470,1) or y.shape == (1470,), "Make sure that y is a single column, Attrition, with all rows"

```

**Instructions**: Now split `X` and `y` into `X_train`, `X_test`, `y_train`, `y_test`.

Relevant documentation:
* [`train_test_split`](https://scikit-learn.org/stable/modules/generated/sklearn.model_selection.train_test_split.html).

<!-- grade_id: cell-54aea9ec522367b7, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import train_test_split

# TODO: Split the data in a 80/20 training/test split, and use the random_state 1234
# YOUR CODE HERE
X_train, X_test, y_train, y_test = train_test_split(
    X, y, test_size=0.2, random_state=1234)
```

```python
print(X_train.shape)
print(X_test.shape)
print(y_train.shape)
print(y_test.shape)
```

Output:
```
(1176, 34)
(294, 34)
(1176,)
(294,)
```

```python

```

<!-- grade_id: cell-7a9828a2155d601c, points: 0, grade: true, solution: false -->
```python
# Automated tests - make sure you pass these!
assert X_train.shape == (1176, 34)
assert X_test.shape == (294, 34)
assert y_train.shape == (1176,)
assert y_test.shape == (294,)

assert X_train.MonthlyRate.iloc[0] == 19146, "Make sure you're using the correct random_state"
```

## Train The Model?

Now it's time to train the k-NN model!

**Instructions**: Create and train a `KNearestNeighbors` model using your training data.

Relevant documentation:
* [`KNeighborsClassifier.fit`](https://scikit-learn.org/stable/modules/generated/sklearn.neighbors.KNeighborsClassifier.html#sklearn.neighbors.KNeighborsClassifier.fit).

<!-- grade_id: cell-1cf4826b8366864f, points: 0, grade: false, solution: true -->
```python
from sklearn.neighbors import KNeighborsClassifier

knn = None
# TODO: Create a k-NN Model with default parameters
# YOUR CODE HERE
neigh = KNeighborsClassifier(n_neighbors=5)
knn = neigh
```

```python
knn = neigh
```

<!-- grade_id: cell-bf4b9b937222fb3f, points: 1, grade: true, solution: false -->
```python
assert type(knn) == KNeighborsClassifier
```

```python
# TODO: Now try fitting the model - do you get an error?

knn = neigh.fit(X, y)
```

Output:
```
c:\Users\13601\OneDrive\Desktop\study\code workspace\522\CSC-522-PJ\myvenv\Lib\site-packages\sklearn\utils\_array_api.py in ?(array, dtype, order, copy, xp, device)
    835         # Use NumPy API to support order
    836         if copy is True:
    837             array = numpy.array(array, order=order, dtype=dtype)
    838         else:
--> 839             array = numpy.asarray(array, order=order, dtype=dtype)
    840 
    841         # At this point array is a NumPy ndarray. We convert it to an array
    842         # container that is consistent with the input's namespace.

c:\Users\13601\OneDrive\Desktop\study\code workspace\522\CSC-522-PJ\myvenv\Lib\site-packages\pandas\core\generic.py in ?(self, dtype, copy)
   2149     def __array__(
   2150         self, dtype: npt.DTypeLike | None = None, copy: bool_t | None = None
   2151     ) -> np.ndarray:
   2152         values = self._values
-> 2153         arr = np.asarray(values, dtype=dtype)
   2154         if (
   2155             
```

## Select Usable Features

You probably just got a value error:
```
ValueError: could not convert string to float: 'Travel_Rarely'
```


**Instructions**: Read the [KNeighborsClassifier](https://scikit-learn.org/stable/modules/generated/sklearn.neighbors.KNeighborsClassifier.html#kneighborsclassifier) documentation figure out what default distance metric it uses, under the `metric` parameter. Write the metric below and explain why it does not work on your dataset.

(**Double-click this to edit!**)

Default distance metric: minkowski

Reason for failure: Travel_Rarely, Travel_Frequently can not be use for calculater Euclidean distance

```python
# The .dtype property gives the data types of each column in our dataset

X_train.dtypes
```

Output:
```
Age                          int64
BusinessTravel              object
DailyRate                    int64
Department                  object
DistanceFromHome             int64
Education                    int64
EducationField              object
EmployeeCount                int64
EmployeeNumber               int64
EnvironmentSatisfaction      int64
Gender                      object
HourlyRate                   int64
JobInvolvement               int64
JobLevel                     int64
JobRole                     object
JobSatisfaction              int64
MaritalStatus               object
MonthlyIncome                int64
MonthlyRate                  int64
NumCompaniesWorked           int64
Over18                      object
OverTime                    object
PercentSalaryHike            int64
PerformanceRating            int64
RelationshipSatisfaction     int64
StandardHours                int64
StockOptionLevel             int64
TotalWorkingYears            int64
TrainingTimesLastYea
```

**Instructions**: To avoid getting the `ValueError`, we need to remove non-numeric attributes. To do this, we should use a ColumnSelector. We will use this so the model knows what columns to keep.

Relevant documentation:
* [`make_column_selector`](https://scikit-learn.org/stable/modules/generated/sklearn.compose.make_column_selector.html): Pay attention to the `dtype_include` parameter.

<!-- grade_id: cell-40caac0a30e9a9ec, points: 0, grade: false, solution: true -->
```python
from sklearn.compose import make_column_selector
import numpy as np

numeric_selector = None
# TODO: Instantiate numeric_selector to include only numeric data types
# YOUR CODE HERE
numeric_selector = make_column_selector(dtype_include=np.number)

# Test numeric_selector - it is a function that takes in a DataFrame
# and returns only desired column names
numeric_selector(X_train)
```

Output:
```
['Age',
 'DailyRate',
 'DistanceFromHome',
 'Education',
 'EmployeeCount',
 'EmployeeNumber',
 'EnvironmentSatisfaction',
 'HourlyRate',
 'JobInvolvement',
 'JobLevel',
 'JobSatisfaction',
 'MonthlyIncome',
 'MonthlyRate',
 'NumCompaniesWorked',
 'PercentSalaryHike',
 'PerformanceRating',
 'RelationshipSatisfaction',
 'StandardHours',
 'StockOptionLevel',
 'TotalWorkingYears',
 'TrainingTimesLastYear',
 'WorkLifeBalance',
 'YearsAtCompany',
 'YearsInCurrentRole',
 'YearsSinceLastPromotion',
 'YearsWithCurrManager']
```

<!-- grade_id: cell-13e8ec2c3453a525, points: 1, grade: true, solution: false -->
```python
assert len(numeric_selector(X_train)) == 26, "numeric_selector(X_train) should return the 26 numeric columns"
```

Now we need a way to drop the non-numeric columns. There are many ways to do this, but we want to use a ColumnTransformer

Relevant documentation:
* [`make_column_transformer`](https://scikit-learn.org/stable/modules/generated/sklearn.compose.make_column_transformer.html)

```python
from sklearn.compose import make_column_transformer

# We create a transformer that takes in our dataset and keeps only the numeric columns (identified by our numeric selector)
numeric_only_transformer = make_column_transformer(
    ('passthrough', numeric_selector)
)
# Notice the shape of the output has changed, containing only 26 columns
numeric_only_transformer.fit_transform(X_train).shape

```

Output:
```
(1176, 26)
```

```python
X_train.shape
```

Output:
```
(1176, 34)
```

```python

```

## Create a Pipeline

```python
# You should have created a knn classifier earlier
knn = KNeighborsClassifier()
knn
```

Output:
```
KNeighborsClassifier()
```

**Instructions**: Finally, we can combine our `numeric_only_transformer` with our `knn` model into a single pipeline.

Pipelines are useful because they can be trained (`fit`) on the training data or used for inference (`predict`) on the test data. 

For our pipeline, in both cases, the data will go through two steps:
1. The non-numeric columns are be dropped
2. The data is passed into the knn model for either training or prediction

While we could have manually dropped non-numeric columns from our dataset, the pipeline approach is powerful because:
* We can easily change what preprocessing steps we are doing, without modifying our original dataset.
* We can apply the pipeline directly to data that has not been preprocessed (e.g. newly collected data).

We will learn more about pipelines in the next milestone.

Relevant documentation:
* [`make_pipeline`](https://scikit-learn.org/stable/modules/generated/sklearn.pipeline.make_pipeline.html)

<!-- grade_id: cell-ffc85d5015ba2abf, points: 0, grade: false, solution: true -->
```python
from sklearn.pipeline import make_pipeline

knn_pipeline = None
# TODO: Use make_pipeline to construct a pipeline
# with numeric_only_transformer and knn
# YOUR CODE HERE
knn_pipeline = make_pipeline(numeric_only_transformer,knn)
knn_pipeline
```

Output:
```
Pipeline(steps=[('columntransformer',
                 ColumnTransformer(transformers=[('passthrough', 'passthrough',
                                                  <sklearn.compose._column_transformer.make_column_selector object at 0x00000273C93D68A0>)])),
                ('kneighborsclassifier', KNeighborsClassifier())])
```

<!-- grade_id: cell-8ae1990db1add2a4, points: 1, grade: true, solution: false -->
```python
assert len(knn_pipeline.steps), "Make sure that you have 2 steps in your pipeline"
assert knn_pipeline.steps[0][0] == "columntransformer", "Make sure the first step is a the numeric_only_transformer"
assert knn_pipeline.steps[1][0] == "kneighborsclassifier", "Make sure the second step is a the knn model"
```

**Instructions**: Train the whole pipeline on the training data. You can use the same `fit` function you tried to use earlier with the knn model.

Relevant documentation:
* [`Pipeline.fit`](https://scikit-learn.org/stable/modules/generated/sklearn.pipeline.Pipeline.html#sklearn.pipeline.Pipeline.fit)

<!-- grade_id: cell-b32ada9ccdc891c5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
knn_pipeline.fit(X_train, y_train)
```

Output:
```
Pipeline(steps=[('columntransformer',
                 ColumnTransformer(transformers=[('passthrough', 'passthrough',
                                                  <sklearn.compose._column_transformer.make_column_selector object at 0x00000273C93D68A0>)])),
                ('kneighborsclassifier', KNeighborsClassifier())])
```

<!-- grade_id: cell-3ec9838730435418, points: 1, grade: true, solution: false -->
```python
from sklearn.utils.validation import check_is_fitted

assert not check_is_fitted(knn_pipeline), "Don't forget to train your pipeline"
```

## Evaluate the Pipeline

**Instructions**: Now let's evaluate our classifier's performance. Let's start by evaluating it on the **training dataset**. It should do well on this dataset, since it's already seen it, but let's make sure.

Relevant documentation:
* [`Pipeline.predict`](https://scikit-learn.org/stable/modules/generated/sklearn.pipeline.Pipeline.html#sklearn.pipeline.Pipeline.predict)

<!-- grade_id: cell-5e8d55898461ff95, points: 0, grade: false, solution: true -->
```python
training_predictions = None
# TODO: Predict the y-values of our X_train dataset.
# YOUR CODE HERE
training_predictions = knn_pipeline.predict(X_train)

# The training predictions should be a vector of Yes/No values, mostly No
training_predictions

```

Output:
```
array(['No', 'No', 'No', ..., 'No', 'No', 'No'],
      shape=(1176,), dtype=object)
```

```python
training_predictions.shape
```

Output:
```
(1176,)
```

```python
y_train
```

Output:
```
630      No
1221     No
122     Yes
706     Yes
87       No
       ... 
1228     No
1077    Yes
1318     No
723      No
815      No
Name: Attrition, Length: 1176, dtype: object
```

```python
# Now let's see how accurate we were.
# We can use vector operations to compare the y' and y vectors
# Taking the mean gets us the average of how often they were equal
(training_predictions == y_train).mean()
```

Output:
```
np.float64(0.8579931972789115)
```

```python
from sklearn.metrics import accuracy_score

# This is equivallent to accuracy, and sklearn has a helper function for that
test_predictions = knn_pipeline.predict(X_test)
accuracy_score(y_test, test_predictions)
```

Output:
```
0.8197278911564626
```

<!-- grade_id: cell-a493e3e5c5854615, points: 0, grade: true, solution: false -->
```python
from numpy.testing import assert_almost_equal

assert_almost_equal((training_predictions == y_train).mean(), 0.8579931972789115)
```

**Instructions**: Now evaluate the model's accuracy on the **test dataset**.

Relevant documentation (you have two options for how to do it):
* [`accuracy_score`](https://scikit-learn.org/stable/modules/generated/sklearn.metrics.accuracy_score.html): Pay close attention to the order of the arguments!
* [`Pipeline.score`](https://scikit-learn.org/1.5/modules/generated/sklearn.pipeline.Pipeline.html#sklearn.pipeline.Pipeline.score): This is a little easier, if you don't need to know that actual predictions.

<!-- grade_id: cell-90c4d9672d240bf6, points: 0, grade: false, solution: true -->
```python
test_accuracy = None
# TODO: Calculate test accuracy
# YOUR CODE HERE
test_accuracy = accuracy_score(y_test, test_predictions)
test_accuracy
```

Output:
```
0.8197278911564626
```

<!-- grade_id: cell-7fa2f3013ba979c3, points: 1, grade: true, solution: false -->
```python
assert_almost_equal(test_accuracy, 0.8197278911564626)
```

This helpfer function returns a DataFrame with the training and evaluation (e.g. test) accuracy for a list of `pipelines`, passed in as the first argument, using a given set of: `X_train`, `y_train`, `X_evaluate`, `y_evaluate`.

**Note**: The input is a list of (name, pipeline) tuples, e.g.
```
evaluate_pipelines(
    [('name1', pipeline_1), ('name2', pipeline_2)],
    X_train, y_train, X_test, y_test
)
```

```python
def evaluate_pipelines(pipelines, X_train, y_train, X_evaluate, y_evaluate):
    rows = []
    for name, pipeline in pipelines:
        pipeline.fit(X_train, y_train)
        training_accuracy = pipeline.score(X_train, y_train)
        test_accuracy = pipeline.score(X_evaluate, y_evaluate)
        rows.append({'name': name, 'training_accuracy': training_accuracy, 'eval_accuracy': test_accuracy})
    return pd.DataFrame(rows)
```

```python
# This makes it easier to get the values we calculated earlier
evaluate_pipelines([('knn', knn_pipeline)], X_train, y_train, X_test, y_test)
```

Output:
```
name  training_accuracy  eval_accuracy
0  knn           0.857993       0.819728
```

## Explore the Effect of `k`

**Instructions**: We know that the k parameter in k-NN controls the complexity of the model. Let's see how well it does with different values of k. You'll analyze the results in your reflection.

Relevant documentation:
* [`KNeighborsClassifier`](https://scikit-learn.org/stable/modules/generated/sklearn.neighbors.KNeighborsClassifier.html): It calls the k parameter `n_neighbors`

<!-- grade_id: cell-c633085fff98c12e, points: 0, grade: false, solution: true -->
```python
pipelines = []
k_values = range(1, 30, 4)

# TODO: Create a list of knn pipelines, each with a different value of k
# given by the k_values variable
for i in k_values:
    knn_pipeline_i = None
    # YOUR CODE HERE
    
    knn = KNeighborsClassifier(n_neighbors=i)
    knn_pipeline_i = make_pipeline(numeric_only_transformer,knn)
    # Append a tuple with the name and pipeline
    pipelines.append((f'knn_{i}', knn_pipeline_i))

evaluate_pipelines(pipelines, X_train, y_train, X_test, y_test)
```

Output:
```
name  training_accuracy  eval_accuracy
0   knn_1           1.000000       0.775510
1   knn_5           0.857993       0.819728
2   knn_9           0.846088       0.836735
3  knn_13           0.840136       0.829932
4  knn_17           0.840986       0.833333
5  knn_21           0.839286       0.836735
6  knn_25           0.838435       0.840136
7  knn_29           0.838435       0.840136
```

<!-- grade_id: cell-9db4633135b59095, points: 2, grade: true, solution: false -->
```python
assert len(pipelines) == len(k_values), f"Make sure you have {len(k_values)} pipelines in your list"
for i in range(i, len(k_values)):
    assert pipelines[i][1].steps[1][1].n_neighbors == k_values[i], f"Make sure your {i}th pipeline has a k-value of {k_values[i]}"
```

Now, let's compare that performance with how a "Majority-class Classifier" baseline would perform by always predicting "No".

Another way of phrasing that is
1. What is the majority class?
2. What % of `y_test` is equal to that class?

```python
majority_percentage = (y_test.value_counts('NO') / len(y_test)) * 100
```

```python
counts = y_test.value_counts() 
y_test.value_counts()
```

Output:
```
Attrition
No     247
Yes     47
Name: count, dtype: int64
```

```python
majority_percentage = (counts['No']/ len(y_test)) * 100
majority_percentage
```

Output:
```
np.float64(84.01360544217688)
```

<!-- grade_id: cell-944d149ca922877c, points: 0, grade: false, solution: true -->
```python
# TODO: Calculate the performance of a majority-class classifier
baseline_test_accuracy = None
# YOUR CODE HERE
baseline_test_accuracy = majority_percentage
baseline_test_accuracy
```

Output:
```
np.float64(84.01360544217688)
```

<!-- grade_id: cell-6207412a8fa3f4e7, points: 2, grade: true, solution: false -->
```python
assert_almost_equal(test_accuracy, 0.8197278911564626)
```

**Remember**: Make sure to complete all problems (.ipynb files) in this assignment. When you finish, double-check the submission instructions at the top of this file, and submit on JupyterHub.